## Introduction
Hello everyone, and welcome to the Introduction to Julia course. While Julia does have a command line interface, we will be using Jupyter (Julia, Python, R) notebooks. Incase you haven't used these before and to familiarize ourselves with them, we will be going over some basic programming concepts before moving on to more advanced topics.

# Getting started

Topics:
1.  How to print
2. How to assign variables
3. How to comment
4. Syntax for basic math

## How to print

In Julia we usually use `println()` to print

```julia
println("I'm excited to learn Julia!")
```

And to access documentation for the println() and other functions


```julia
?println
```

## How to assign variables

All we need is a variable name, value, and an equal's sign!<br>
Julia is a dynamic language and thus will figure out types for us.

```julia
my_answer = 42
typeof(my_answer)
```

```julia
my_pi = 3.14159
typeof(my_pi)
```

```julia
😺 = "smiley cat!"
typeof(😺)
```

To type a smiley cat, use tab completion to select the emoji name and then tab again

```julia
# \:smi + <tab> --> select with down arrow + <enter> ---> <tab> + <enter> to complete
```

After assigning a value to a variable, we can reassign a value of a different type to that variable without any issue.

```julia
😺 = 1
```

```julia
typeof(😺)

```

## How to comment

```julia
# You can leave comments on a single line using the pound/hash key
```

```julia
#=

For multi-line comments, 
use the '#= =#' sequence.

=#
```

## Syntax for basic math

```julia

```

```julia
sum = 3 + 7
```

```julia
difference = 10 - 3
```

```julia
product = 20 * 5
```

```julia
quotient = 100 / 10
```

```julia
power = 10 ^ 2
```

```julia
modulus = 101 % 2
```

# Numbers
Many different ways to write the number forty-two


```julia
fortytwos = (42, 42.0, 4.20e1, 4.20f1, 84//2, 0x2a)
```

```julia
for x in fortytwos
    show(x)
    println("\tisa $(typeof(x))")
end
```

# Updating Operators
All the infix arithmetic operators above can be used as "updating" operators in conjunction with an assignment:

```julia
y = 1
y += 1

```

```julia
println(y)

```

```julia
y = 1
y = y + 1

```

```julia
println(y)
```

Note that it's just re-purposing the same name for a new value. This means that the type can change!

```julia
y /= 2
```

```julia
typeof(y)
```

# Strings and characters

```julia
s1 = "hello world"

```

```julia
s2 = """ hello world"

```

```julia
s3 = "Here, we get an "error" because it's ambiguous where this string ends "
```

```julia
s4 = """ no "errors" here """
```

```julia
println("""The other nice thing about triple-quoted
               string literals is that they ignore leading
               indentation, which is nice for long strings
               in real code. Try changing these quotes!""")


```

Strings are not written with single 's — that's used for a single character:

```julia
first(s1)
```

```julia
c1 = 's'
```

```julia
'If you try writing a string in single-quotes, you will get an error'
```

# String interpolation


You can use the dollar sign inside a string to evaluate a Julia expression inside a string — either a single variable or a more complicated expression:


```julia
name = "Yiannis"
num_fingers = 10
num_toes = 10 
println("Hello, my name is $name.")
println("I have $num_fingers fingers and $num_toes toes.")
```

What do you think the following will print?

```julia
println("That is $(num_fingers+num_toes) digits in all")
```

### Exercises

#### 1.1 
Look up docs for the `convert` function.

```julia

```

#### 1.2 
Assign `365` to a variable named `days`. Convert `days` to a float and assign it to variable `days_float`
If you've done this correctly, the @assert function will check it for you.

```julia

```

<!-- grade_id: cell-715f78016beb0489, points: 1, grade: true, solution: false -->
```julia
@assert days == 365
@assert days_float == 365.0

```

#### 1.3 
See what happens when you execute

```julia
convert(Int64, "1")
```
and

```julia
parse(Int64, "1")
```

```julia

```

```julia

```