# Data structures

Once we start working with many pieces of data at once, it will be convenient for us to store data in structures like arrays or dictionaries (rather than just relying on variables).

Types of data structures covered:
1. Tuples
2. Dictionaries
3. Arrays


As an overview, tuples and arrays are both ordered sequences of elements (so we can index into them). 

Dictionaries and arrays are both **mutable**, whereas tuples are **immutable**

We'll explain this more below!
Arrays (in my opinion) are one of the most important structures in (Scientific) Computing.

## Tuples

We can create a tuple by enclosing an ordered collection of elements in `( )`.

Syntax: <br>
```julia
(item1, item2, ...)
```


```julia
myfavoriteanimals = ("penguins", "cats", "sugargliders")
```

We can index into this tuple,

```julia
myfavoriteanimals[1]
```

but since tuples are immutable, we can't update it

```julia
myfavoriteanimals[1] = "otters"
```

## Julia Specific: NamedTuples

As you might guess, `NamedTuple`s are just like `Tuple`s except that each element additionally has a name! They have a special syntax using `=` inside a tuple:

```julia
(name1 = item1, name2 = item2, ...)
```

```julia
myfavoriteanimals = (bird = "penguins", mammal = "cats", marsupial = "sugargliders")
```

Like regular `Tuples`, `NamedTuples` are ordered, so we can retrieve their elements via indexing:

```julia
myfavoriteanimals[1]
```

They also add the special ability to access values by their name:

```julia
myfavoriteanimals.bird
```

## Dictionaries

If we have sets of data related to one another, we may choose to store that data in a dictionary. We can create a dictionary using the `Dict()` function, which we can initialize as an empty dictionary or one storing key, value pairs.

Syntax:
```julia
Dict(key1 => value1, key2 => value2, ...)
```

A good example is a contacts list, where we associate names with phone numbers.

```julia
myphonebook = Dict("Yiannis" => "0044000000", "Imperial" => "02075895111")
```

In this example, each name and number is a "key" and "value" pair. We can grab Jenny's number (a value) using the associated key

```julia
myphonebook["Yiannis"]
```

We can add another entry to this dictionary as follows

```julia
myphonebook["Graduate School"] = "020 7589 5111"
```

Let's check what our phonebook looks like now...

```julia
myphonebook
```

We can delete Yiannis from our contact list - and simultaneously grab his number - by using `pop!`

```julia
pop!(myphonebook, "Yiannis")
```

```julia
myphonebook
```

Unlike tuples and arrays, dictionaries are not ordered. So, we can't index into them.

```julia
myphonebook[1]
```

In the example above, `julia` thinks you're trying to access a value associated with the key `1`.

```julia

```

## Arrays

Unlike tuples, arrays are mutable. Unlike dictionaries, arrays contain ordered collections. <br>
We can create an array by enclosing this collection in `[ ]`.

Syntax: <br>
```julia
[item1, item2, ...]
```

For an interesting piece of Julia history regarding differences between matrices and vectors feel free to read https://github.com/julialang/julia/issues/4774


For example, we might create an array to keep track of pokemon


```julia
mypokemon = ["Bulbasaur","Squirtle","Charmander","Pikachu"]
```

The `1` in `Array{String,1}` means this is a one dimensional vector.  An `Array{String,2}` would be a 2d matrix, etc.  The `String` is the type of each element.

or to store a sequence of numbers

```julia
fibonacci = [1, 1, 2, 3, 5, 8, 13]
```

```julia

```

```julia
mixture = [1, 1, 2, 3, "Yiannis", "Tony"]
```

Once we have an array, we can grab individual pieces of data from inside that array by indexing into the array. For example, if we want the third pokemon listed in `pokemon`, we write

```julia
pokemon[3]
```

We can use indexing to edit an existing element of an array

```julia
pokemon[3] = "Mewtwo"
```

Yes, Julia is 1-based indexing, not 0-based like Python.  (Although https://docs.julialang.org/en/v1/devdocs/offset-arrays/)

We can also edit the array by using the `push!` and `pop!` functions. `push!` adds an element to the end of an array and `pop!` removes the last element of an array. 

Remember the Julia notation !means the object changes.

We can add another number to our fibonnaci sequence

```julia
push!(fibonacci, 21)
```

and then remove it

```julia
pop!(fibonacci)
```

```julia
fibonacci
```

So far I've given examples of only 1D arrays of scalars, but arrays can have an arbitrary number of dimensions and can also store other arrays. 
<br><br>
For example, the following are arrays of arrays:

```julia
favorites = [["koobideh", "chocolate", "eggs"],["penguins", "cats", "sugargliders"]]
```

```julia
numbers = [[1, 2, 3], [4, 5], [6, 7, 8, 9]]
```

Below are examples of 2D and 3D arrays populated with random values.

```julia
rand(4, 3)
```

```julia
rand(4, 3, 2)
```

Be careful when you want to copy arrays!

```julia
fibonacci
```

```julia
somenumbers = fibonacci
```

```julia
somenumbers[1] = 404
```

```julia
fibonacci
```

Editing `somenumbers` caused `fibonacci` to get updated as well!

In the above example, we didn't actually make a copy of `fibonacci`. We just created a new way to access the entries in the array bound to `fibonacci`.

If we'd like to make a copy of the array bound to `fibonacci`, we can use the `copy` function.

```julia
# First, restore fibonacci
fibonacci[1] = 1
fibonacci
```

```julia
somemorenumbers = copy(fibonacci)
```

```julia
somemorenumbers[1] = 404
```

```julia
fibonacci
```

In this last example, fibonacci was not updated. Therefore we see that the arrays bound to `somemorenumbers` and `fibonacci` are distinct.

# Vectors and Matrices
We can also create vectors : 

```julia
squares = [1, 4, 9, 16, 25, 36, 49, 64]
```

```julia
squares[1]
```

```julia
squares[1:3]
```

```julia
squares[end]
```

The below code will give us a first glimpse into the powerful **type** system present in Julia

```julia
typeof(squares) 
```

# Concatenation
If, instead of commas, we use spaces, then the values are concatenated horizontally

```julia
cubes = [1, 8, 27, 64, 125, 216, 343, 512]
```

```julia
powers = [1:8 squares cubes]
```

```julia
powers[4, 2]
```

```julia
powers[:, 3]
```

```julia
powers[7, :]
```

```julia
typeof(powers)
```

Semicolon separators perform vertical concatenation:

```julia
[squares; cubes]
```

Whereas commas would simply create an array of arrays : 

```julia
nested_powers = [[1,2,3,4,5,6,7,8], squares, cubes]
```

```julia
nested_powers[2]
```

Horizontal and vertical concatenation can be used together to as a simple syntax for matrix literals

```julia
[1 3 5; 2 4 6]
```

## for loops

The syntax for a `for` loop is

```julia
for *var* in *loop iterable*
    *loop body*
end
```

## while loops

The syntax for a `while` is

```julia
while *condition*
    *loop body*
end
```

We could use a for loop to generate the same results as either of the examples above:

```julia
A = fill(0, (8, 3)) # Allocate an 8x3 matrix to store the values into
for pow in 1:3
    for value in 1:8
        A[value, pow] = value ^ pow
    end
end
A
```

```julia
A == powers
```


## Array Comprehensions

```julia
squares = [value^2 for value in 1:8]
```

```julia
cubes = [value^3 for value in 1:8]
```

```julia
powers = [value^pow for value in 1:8, pow in 1:3]
```

# The element type

Note that every time an array prints out, it is displaying its element type and dimensionality, for example `Array{Int64, 2}`. This describes what it can store — and thus what it can return upon indexing.

```julia
typeof(powers)
```

```julia
typeof(powers[1, 1])
```

Further, the array will try to convert any new values assigned into it to its element type:

```julia
powers[1, 1] = 1.6
```

```julia
powers[1, 1] = -5.0 # This can be losslessly converted to an integer
```

```julia
powers
```

Arrays that have an exact and concrete element type are generally significantly faster, so Julia will try to find an amenable element type for you in its literal construction syntax:

```julia
fortytwosarray = [42, 42.0, 4.20e1, 4.20f1, 84//2, 0x2a]
```

```julia
for x in fortytwosarray
    show(x)
    println("\tisa $(typeof(x))")
end
```

The `Any` array can be helpful for disabling these behaviors and allowing all kinds of different objects:

```julia
anyfortytwos = Any[42, 42.0, 4.20e1, 4.20f1, 84//2, 0x2a]
```

```julia
anyfortytwos[1] = "FORTY TWO"
anyfortytwos
```

### Exercises

#### 3.1 
Create an array, `a_ray`, with the following code:

```julia
a_ray = [1, 2, 3]
```

Add the number `4` to the end of this array and then remove it.

```julia

```

<!-- grade_id: cell-a6b20b78fbfeb449, points: 1, grade: true, solution: false -->
```julia
@assert a_ray == [1, 2, 3]
```

#### 3.2 
Try to add "Emergency" as key to `myphonebook` with the value `string(911)` with the following code
```julia
myphonebook["Emergency"] = 911
```

Why doesn't this work?

```julia

```

#### 3.3 
Create a new dictionary called `flexible_phonebook` that has Jenny's number stored as an integer and Ghostbusters' number stored as a string with the following code

```julia
flexible_phonebook = Dict("Jenny" => 8675309, "Ghostbusters" => "555-2368")
```

```julia

```

<!-- grade_id: cell-66f1d636c676c2f2, points: 1, grade: true, solution: false -->
```julia
@assert flexible_phonebook == Dict("Jenny" => 8675309, "Ghostbusters" => "555-2368")
```

#### 3.4 
Add the key "Emergency" with the value `911` (an integer) to `flexible_phonebook`.

```julia

```

<!-- grade_id: cell-ee862ab938680259, points: 1, grade: true, solution: false -->
```julia
@assert haskey(flexible_phonebook, "Emergency")
```

<!-- grade_id: cell-c8c40ababe7cdfce, points: 2, grade: true, solution: false -->
```julia
@assert flexible_phonebook["Emergency"] == 911
```

#### 3.5 
Why can we add an integer as a value to `flexible_phonebook` but not `myphonebook`? How could we have initialized `myphonebook` so that it would accept integers as values? (hint: try using [Julia's documentation for dictionaries](https://docs.julialang.org/en/v1/base/collections/#Dictionaries))

#### 3.6 
Use an array comprehension to create an an array `squares_arr` that stores the squares for all integers between 1 and 100.

```julia

```

```julia
@assert length(squares_arr) == 100
@assert sum(squares_arr) == 338350
```