# Factorizations and other fun
Based on work by Andreas Noack Jensen (MIT) (http://www.econ.ku.dk/phdstudent/noack/)

## Outline
 - Factorizations
 - Special matrix structures
 - Generic linear algebra

Before we get started, let's set up a linear system and use `LinearAlgebra` to bring in the factorizations and special matrix structures.

```julia
using LinearAlgebra
A = rand(3, 3)
x = fill(1, (3,))
b = A * x
```

Output:
```
3-element Vector{Float64}:
 1.6170120247910962
 1.6644545921552723
 0.5894827313316517
```

## Factorizations

#### LU factorizations
In Julia we can perform an LU factorization
```julia
PA = LU
``` 
where `P` is a permutation matrix, `L` is lower triangular unit diagonal and `U` is upper triangular, using `lufact`.

Julia allows computing the LU factorization and defines a composite factorization type for storing it.

```julia
Alu = lu(A)
```

Output:
```
LU{Float64, Matrix{Float64}}
L factor:
3×3 Matrix{Float64}:
 1.0       0.0       0.0
 0.124045  1.0       0.0
 0.127652  0.281901  1.0
U factor:
3×3 Matrix{Float64}:
 0.302203  0.42203    0.940222
 0.0       0.694737   0.715808
 0.0       0.0       -0.0206226
```

```julia
?lu
```

Output:
```
search: [0m[1ml[22m[0m[1mu[22m [0m[1ml[22m[0m[1mu[22m! [0m[1mL[22m[0m[1mU[22m A[0m[1ml[22m[0m[1mu[22m f[0m[1ml[22m[0m[1mu[22msh va[0m[1ml[22m[0m[1mu[22mes inc[0m[1ml[22m[0m[1mu[22mde inc[0m[1ml[22m[0m[1mu[22mde_string inc[0m[1ml[22m[0m[1mu[22mde_dependency

[36m  lu(A, pivot=Val(true); check = true) -> F::LU[39m

  Compute the LU factorization of [36mA[39m.

  When [36mcheck = true[39m, an error is thrown if the decomposition fails. When [36mcheck
  = false[39m, responsibility for checking the decomposition's validity (via
  [36missuccess[39m) lies with the user.

  In most cases, if [36mA[39m is a subtype [36mS[39m of [36mAbstractMatrix{T}[39m with an element type
  [36mT[39m supporting [36m+[39m, [36m-[39m, [36m*[39m and [36m/[39m, the return type is [36mLU{T,S{T}}[39m. If pivoting is
  chosen (default) the element type should also support [36mabs[39m and [36m<[39m.

  The individual components of t
```

```julia
typeof(Alu)
```

Output:
```
LU{Float64, Matrix{Float64}}
```

The different parts of the factorization can be extracted by accessing their special properties

```julia
Alu.P
```

Output:
```
3×3 Matrix{Float64}:
 0.0  1.0  0.0
 1.0  0.0  0.0
 0.0  0.0  1.0
```

```julia
Alu.L
```

Output:
```
3×3 Matrix{Float64}:
 1.0       0.0       0.0
 0.124045  1.0       0.0
 0.127652  0.281901  1.0
```

```julia
Alu.U
```

Output:
```
3×3 Matrix{Float64}:
 0.302203  0.42203    0.940222
 0.0       0.694737   0.715808
 0.0       0.0       -0.0206226
```

Julia can dispatch methods on factorization objects.

For example, we can solve the linear system using either the original matrix or the factorization object.

```julia
A\b
```

Output:
```
3-element Vector{Float64}:
 1.0
 1.0
 1.0
```

```julia
Alu\b
```

Output:
```
3-element Vector{Float64}:
 1.0
 1.0
 1.0
```

Similarly, we can calculate the determinant of `A` using either `A` or the factorization object

```julia
det(A) ≈ det(Alu)
```

Output:
```
true
```

#### QR factorizations

In Julia we can perform a QR factorization
```
A=QR
``` 

where `Q` is unitary/orthogonal and `R` is upper triangular, using `qrfact`. 

```julia
Aqr = qr(A)
```

Output:
```
LinearAlgebra.QRCompactWY{Float64, Matrix{Float64}}
Q factor:
3×3 LinearAlgebra.QRCompactWYQ{Float64, Matrix{Float64}}:
 -0.122125   0.955012  -0.270253
 -0.984526  -0.151044  -0.0888543
 -0.125677   0.25522    0.95868
R factor:
3×3 Matrix{Float64}:
 -0.306952  -0.538122  -1.06519
  0.0        0.713466   0.729843
  0.0        0.0       -0.0197705
```

```julia
?qr
```

Output:
```
search: [0m[1mq[22m[0m[1mr[22m [0m[1mq[22m[0m[1mr[22m! [0m[1mQ[22m[0m[1mR[22m [0m[1mQ[22m[0m[1mR[22mPivoted A[0m[1mq[22m[0m[1mr[22m s[0m[1mq[22m[0m[1mr[22mt is[0m[1mq[22m[0m[1mr[22mt [0m[1mQ[22muickSo[0m[1mr[22mt Partial[0m[1mQ[22muickSo[0m[1mr[22mt

[36m  qr(A, pivot=Val(false); blocksize) -> F[39m

  Compute the QR factorization of the matrix [36mA[39m: an orthogonal (or unitary if [36mA[39m
  is complex-valued) matrix [36mQ[39m, and an upper triangular matrix [36mR[39m such that

[35m  A = Q R[39m

  The returned object [36mF[39m stores the factorization in a packed format:

    •  if [36mpivot == Val(true)[39m then [36mF[39m is a [36mQRPivoted[39m object,

    •  otherwise if the element type of [36mA[39m is a BLAS type ([36mFloat32[39m,
       [36mFloat64[39m, [36mComplexF32[39m or [36mComplexF64[39m), then [36mF[39m is a [36mQRCompactWY[39m
       object,

    •  otherwise [36mF[39m is a [3
```

Similarly to the LU factorization, the matrices `Q` and `R` can be extracted from the QR factorization object via

```julia
Aqr.Q
```

Output:
```
3×3 LinearAlgebra.QRCompactWYQ{Float64, Matrix{Float64}}:
 -0.122125   0.955012  -0.270253
 -0.984526  -0.151044  -0.0888543
 -0.125677   0.25522    0.95868
```

```julia
Aqr.R
```

Output:
```
3×3 Matrix{Float64}:
 -0.306952  -0.538122  -1.06519
  0.0        0.713466   0.729843
  0.0        0.0       -0.0197705
```

#### Eigendecompositions

The results from eigendecompositions, singular value decompositions, Hessenberg factorizations, and Schur decompositions are all stored in `Factorization` types.

The eigendecomposition can be computed

```julia
Asym = A + A'
AsymEig = eigen(Asym)
```

Output:
```
Eigen{Float64, Float64, Matrix{Float64}, Vector{Float64}}
values:
3-element Vector{Float64}:
 -0.6616442806377435
 -0.4595332606156878
  2.6425825943260914
vectors:
3×3 Matrix{Float64}:
  0.854193  -0.224749  -0.468873
 -0.508131  -0.552044  -0.661097
 -0.110258   0.802954  -0.585755
```

The values and the vectors can be extracted from the Eigen type by special indexing

```julia
AsymEig.values
```

Output:
```
3-element Vector{Float64}:
 -0.6616442806377435
 -0.4595332606156878
  2.6425825943260914
```

```julia
AsymEig.vectors
```

Output:
```
3×3 Matrix{Float64}:
  0.854193  -0.224749  -0.468873
 -0.508131  -0.552044  -0.661097
 -0.110258   0.802954  -0.585755
```

Once again, when the factorization is stored in a type, we can dispatch on it and write specialized methods that exploit the properties of the factorization, e.g. that $A^{-1}=(V\Lambda V^{-1})^{-1}=V\Lambda^{-1}V^{-1}$.

```julia
inv(AsymEig)*Asym
```

Output:
```
3×3 Matrix{Float64}:
  1.0          9.99201e-16  9.4369e-16
  6.66134e-16  1.0          1.22125e-15
 -2.22045e-16  4.44089e-16  1.0
```

## Special matrix structures
Matrix structure is very important in linear algebra. To see *how* important it is, let's work with a larger linear system

```julia
n = 1000
A = randn(n,n);
```

Julia can often infer special matrix structure

```julia
Asym = A + A'
issymmetric(Asym)
```

Output:
```
true
```

but sometimes floating point error might get in the way.

```julia
Asym_noisy = copy(Asym)
Asym_noisy[1,2] += 5eps()
```

Output:
```
-1.7806400370270028
```

```julia
issymmetric(Asym_noisy)
```

Output:
```
false
```

Luckily we can declare structure explicitly with, for example, `Diagonal`, `Triangular`, `Symmetric`, `Hermitian`, `Tridiagonal` and `SymTridiagonal`.

```julia
Asym_explicit = Symmetric(Asym_noisy);
```

Let's compare how long it takes Julia to compute the eigenvalues of `Asym`, `Asym_noisy`, and `Asym_explicit`

```julia
@time eigvals(Asym);
```

Output:
```
0.460447 seconds (133.94 k allocations: 15.454 MiB, 3.51% gc time, 36.15% compilation time)
```

```julia
@time eigvals(Asym_noisy);
```

Output:
```
1.064262 seconds (13 allocations: 7.920 MiB)
```

```julia
@time eigvals(Asym_explicit);
```

Output:
```
0.293277 seconds (5.88 k allocations: 8.358 MiB, 3.56% compilation time)
```

```julia
using BenchmarkTools  

```

```julia
@benchmark eigvals($Asym)
```

Output:
```
BenchmarkTools.Trial: 
  memory estimate:  7.99 MiB
  allocs estimate:  11
  --------------
  minimum time:     265.792 ms (0.00% GC)
  median time:      275.191 ms (0.00% GC)
  mean time:        274.591 ms (0.31% GC)
  maximum time:     283.508 ms (0.00% GC)
  --------------
  samples:          19
  evals/sample:     1
```

```julia
@benchmark eigvals($Asym_noisy)
```

Output:
```
BenchmarkTools.Trial: 
  memory estimate:  7.92 MiB
  allocs estimate:  13
  --------------
  minimum time:     1.049 s (0.00% GC)
  median time:      1.056 s (0.00% GC)
  mean time:        1.068 s (0.00% GC)
  maximum time:     1.112 s (0.00% GC)
  --------------
  samples:          5
  evals/sample:     1
```

```julia
@benchmark eigvals($Asym_explicit)
```

Output:
```
BenchmarkTools.Trial: 
  memory estimate:  7.99 MiB
  allocs estimate:  11
  --------------
  minimum time:     264.836 ms (0.00% GC)
  median time:      271.155 ms (0.00% GC)
  mean time:        271.757 ms (0.26% GC)
  maximum time:     278.987 ms (0.00% GC)
  --------------
  samples:          19
  evals/sample:     1
```

In this example, using `Symmetric()` on `Asym_noisy` made our calculations about `5x` more efficient :)

```julia

```

#### A big problem
Using the `Tridiagonal` and `SymTridiagonal` types to store tridiagonal matrices makes it possible to work with potentially very large tridiagonal problems. The following problem would not be possible to solve on a laptop if the matrix had to be stored as a (dense) `Matrix` type.

```julia
n = 1_000_000;
A = SymTridiagonal(randn(n), randn(n-1));
@time eigmax(A)
```

Output:
```
0.819444 seconds (782.01 k allocations: 226.927 MiB, 2.45% gc time, 28.81% compilation time)
6.19877544505124
```

## Generic linear algebra
The usual way of adding support for numerical linear algebra is by wrapping BLAS and LAPACK subroutines. For matrices with elements of `Float32`, `Float64`, `Complex{Float32}` or `Complex{Float64}` this is also what Julia does.

However, Julia also supports generic linear algebra, allowing you to, for example, work with matrices and vectors of rational numbers.

#### Rational numbers
Julia has rational numbers built in. To construct a rational number, use double forward slashes:

```julia
1//2
```

Output:
```
1//2
```

#### Example: Rational linear system of equations
The following example shows how linear system of equations with rational elements can be solved without promoting to floating point element types. Overflow can easily become a problem when working with rational numbers so we use `BigInt`s.

```julia
Arational = Matrix{Rational{BigInt}}(rand(1:10, 3, 3))/10
```

Output:
```
3×3 Matrix{Rational{BigInt}}:
 2//5  1//1  2//5
 1//1  1//1  7//10
 2//5  4//5  2//5
```

```julia
x = fill(1, 3)
b = Arational*x
```

Output:
```
3-element Vector{Rational{BigInt}}:
  9//5
 27//10
  8//5
```

```julia
Arational\b
```

Output:
```
3-element Vector{Rational{BigInt}}:
 1//1
 1//1
 1//1
```

```julia
lu(Arational)
```

Output:
```
LU{Rational{BigInt}, Matrix{Rational{BigInt}}}
L factor:
3×3 Matrix{Rational{BigInt}}:
 1//1  0//1  0//1
 2//5  1//1  0//1
 2//5  2//3  1//1
U factor:
3×3 Matrix{Rational{BigInt}}:
 1//1  1//1  7//10
 0//1  3//5  3//25
 0//1  0//1  1//25
```

### Exercises

#### 11.1
What are the eigenvalues of matrix A?

```
A =
[
 140   97   74  168  131
  97  106   89  131   36
  74   89  152  144   71
 168  131  144   54  142
 131   36   71  142   36
]
```
and assign it a variable `A_eigv`

```julia

```

```julia
A =
[
 140   97   74  168  131
  97  106   89  131   36
  74   89  152  144   71
 168  131  144   54  142
 131   36   71  142   36
]
```

Output:
```
5×5 Matrix{Int64}:
 140   97   74  168  131
  97  106   89  131   36
  74   89  152  144   71
 168  131  144   54  142
 131   36   71  142   36
```

```julia

```

Output:
```
Eigen{Float64, Float64, Matrix{Float64}, Vector{Float64}}
values:
5-element Vector{Float64}:
 -128.49322764802145
  -55.887784553057
   42.752167279318854
   87.16111477514494
  542.4677301466137
vectors:
5×5 Matrix{Float64}:
 -0.147575  -0.647178  -0.010882    0.548903   -0.507907
 -0.256795   0.173068  -0.834628   -0.239864   -0.387253
 -0.185537  -0.239762   0.422161   -0.731925   -0.440631
  0.819704   0.247506   0.0273194   0.0366447  -0.514526
 -0.453805   0.657619   0.352577    0.322668   -0.364928
```

```julia

```

Output:
```
5-element Vector{Float64}:
 -128.49322764802145
  -55.887784553057
   42.752167279318854
   87.16111477514494
  542.4677301466137
```

```julia

```

<!-- grade_id: cell-4d5f60c8a814c789, points: 0, grade: true, solution: false -->
```julia
@assert A_eigv ≈ [-128.49322764802145, -55.887784553056875, 42.7521672793189, 87.16111477514521, 542.4677301466143]
```

#### 11.2 
Create a `Diagonal` matrix from the eigenvalues of `A`.

```julia

```

```julia

```

Output:
```
5×5 Diagonal{Float64, Vector{Float64}}:
 -128.493     ⋅        ⋅        ⋅         ⋅ 
     ⋅     -55.8878    ⋅        ⋅         ⋅ 
     ⋅        ⋅      42.7522    ⋅         ⋅ 
     ⋅        ⋅        ⋅      87.1611     ⋅ 
     ⋅        ⋅        ⋅        ⋅      542.468
```

<!-- grade_id: cell-3b000a3710c9c263, points: 1, grade: true, solution: false -->
```julia
@assert isapprox(A_diag, [-128.493    0.0      0.0      0.0       0.0;
    0.0    -55.8878   0.0      0.0       0.0;
    0.0      0.0     42.7522   0.0       0.0;
    0.0      0.0      0.0     87.1611    0.0;
    0.0 0.0      0.0      0.0     542.468], atol = 0.001)
```

#### 11.3 
Create a `LowerTriangular` matrix from `A` and store it in `A_lowertri`

```julia

```

Output:
```
5×5 LowerTriangular{Int64, Matrix{Int64}}:
 140    ⋅    ⋅    ⋅   ⋅
  97  106    ⋅    ⋅   ⋅
  74   89  152    ⋅   ⋅
 168  131  144   54   ⋅
 131   36   71  142  36
```

<!-- grade_id: cell-b76cee2b4a8777da, points: 0, grade: true, solution: false -->
```julia
@assert A_lowertri ==  [140    0    0    0   0;
  97  106    0    0   0;
  74   89  152    0   0;
 168  131  144   54   0;
 131   36   71  142  36]
```

```julia

```